---
title: "Exercises"
---

## Overview

Exercises are interactive R code chunks that allow readers to directly execute R code and see its results. There are many options associated with tutorial exercises (all of which are described in more detail below):

<table>
<thead>
<tr class="header">
<th>Option</th>
<th>Description</th>
</tr>
</thead>
<tbody>
<tr class="odd">
<td><code>exercise.cap</code></td>
<td>Caption for exercise chunk (defaults to "Code")</td>
</tr>
<tr class="even">
<td><code>exercise.cap.engine</code></td>
<td>Caption for exercise chunk engine (defaults to the chunk engine value)</td>
</tr>
<tr class="odd">
<td><code>exercise.eval</code></td>
<td>Whether to pre-evaluate the exercise so the reader can see some default output (defaults to <code>FALSE</code>).</td>
</tr>
<tr class="even">
<td><code>exercise.lines</code></td>
<td>Lines of code for exercise editor (default to size of code chunk).</td>
</tr>
<tr class="odd">
<td><code>exercise.timelimit</code></td>
<td>Number of seconds to limit execution time to (defaults to 30).</td>
</tr>
<tr class="even">
<td><code>exercise.checker</code></td>
<td>Function used to check exercise answers (e.g., <code>gradethis::grade_learnr()</code>).</td>
</tr>
<tr class="odd">
<td><code>exercise.error.check.code</code></td>
<td>A string containing R code to use for checking code when an exercise evaluation error occurs (e.g., <code>"gradethis::grade_code()"</code>).</td>
</tr>
<tr class="even">
<td><code>exercise.completion</code></td>
<td>Whether to enable code completion in the exercise editor.</td>
</tr>
<tr class="odd">
<td><code>exercise.diagnostics</code></td>
<td>Whether to enable code diagnostics in the exercise editor.</td>
</tr>
<tr class="even">
<td><code>exercise.startover</code></td>
<td>Whether to include a "Start Over" button for the exercise.</td>
</tr>
<tr class="odd">
<td><code>exercise.warn_invisible</code></td>
<td>Whether to display an invisible result warning if the last value returned is invisible.</td>
</tr>
</tbody>
</table>

Note that these options can all be specified either globally or per-chunk. For example, the following code sets global default options using the `setup` chunk and also sets some local options on the `addition` chunk:

<div id="exerciseoptions"></div>
<script type="text/javascript">loadSnippet('exerciseoptions')</script>

There are also some other specialized chunks that can be used with an exercise chunk, including:

1. Exercise [setup chunks](#exercise-setup), which enable you to execute code to setup the environment immediately prior to executing submitted code.

2. Exercise [solution chunks](#exercise-solutions) which enable you to provide a solution to the exercise that can be optionally viewed by users of the tutorial.

The use of these special chunks is also described in detail below.


## Exercise Evaluation

By default, exercise code chunks are NOT pre-evaluated (i.e there is no initial output displayed for them). However, in some cases you may want to show initial exercise output (especially for exercises like the ones above where the user is asked to modify code rather than write new code from scratch).

You can arrange for an exercise to be pre-evaluated (and its output shown) using the `exercise.eval` chunk option. This option can also be set either globally or per-chunk:

<div id="exerciseeval"></div>
<script type="text/javascript">loadSnippet('exerciseeval')</script>

## Exercise Setup

Code chunks with `exercise=TRUE` are evaluated within standalone environments. This means that they don't have access to previous computations from within the document. This constraint is imposed so that users can execute exercises in any order (i.e. correct execution of one exercise never depends on completion of a prior exercise).

You can however arrange for setup code to be run before evaluation of an exercise to ensure that the environment is primed correctly. There are three ways to provide setup code for an exercise:

1. Add code to the global `setup` chunk. This code is run once at the startup of the tutorial and is shared by all exercises within the tutorial. For example:

    <div id="exercisesetup" style="width: 95%;"></div>
    <script type="text/javascript">loadSnippet('exercisesetup')</script>

2. Create a setup chunk that's shared by several exercises. If you don't want to rely on global setup but would rather create setup code that's used by only a handful of exercises you can use the `exercise.setup` chunk attribute to provide the label of another chunk that will perform setup tasks. To illustrate, we'll re-write the previous example to use a shared setup chunk named `prepare-flights`:

    <div id="exercisesetupshared" style="width: 95%;"></div>
    <script type="text/javascript">loadSnippet('exercisesetupshared')</script>

3. Create a setup chunk that's specific to another chunk using a `-setup` chunk suffix. To do this give your exercise chunk a label (e.g. `filter`) then add another chunk with the same label plus a `-setup` suffix (e.g. `filter-setup`). For example:

    <div id="exercisesetupsuffix" style="width: 95%;"></div>
    <script type="text/javascript">loadSnippet('exercisesetupsuffix')</script>

4. If you want to chain setup chunks such that a setup chunk can inherit its parent setup chunk and so on, you can use `exercise.setup` (**Note**: You must use `exercise.setup` for chaining. You cannot rely on the `-setup` suffix labelling scheme.) `learnr` will keep following the trail of `exercise.setup` chunks until there are no more chunks to be found. To demonstrate, let's turn the first exercise in the previous examples to be another setup chunk called `filtered-flights` with its `exercise.setup=prepare-flights`. This will now filter the data and store it and can be referenced inside the `arrange` exercise:

    <div id="exercisesetupchained" style="width: 95%;"></div>
    <script type="text/javascript">loadSnippet('exercisesetupchained')</script>
    
5. You can chain exercise chunks as well, but keep in mind that pre-filled code is used to serve as the setup code instead of user input. For example, if you turned `filtered-flights` back to an exercise, the pre-filled code is used as setup for the `arrange` exercise that use it as its setup:

    <div id="exercisechained" style="width: 95%;"></div>
    <script type="text/javascript">loadSnippet('exercisechained')</script>

## Hints and Solutions

You can optionally provide a hint or solution for each exercise that can be optionally displayed by users. Hints can be based on either R code snippets or on custom markdown/HTML content.

### R Code Hints

To create a hint or solution based on R code simply create a new code chunk with "-hint" or "-solution" chunk label suffix. For example:

<div id="exercisesolution"></div>
<script type="text/javascript">loadSnippet('exercisesolution')</script>

A "Hint" or "Solution" button is added to the left side of the exercise header region:

<img src="images/solution.png"  width="732" height="183"/>

### Markdown Hints

To create a hint based on custom markdown content simply add a `<div>` tag with an `id` attribue that marks it as hint for your exercise (e.g. "filter-hint"). For example:

<div id="exercisehintdiv"></div>
<script type="text/javascript">loadSnippet('exercisehintdiv', 'markdown')</script>

The content within the `<div>` will be displayed underneath the R code editor for the exercise whenever the user presses the "Hint" button.

### Multiple Hints

For R code hints you can provide a sequence of hints that reveal progressively more of the solution as desired by the user. To do this create a sequence of indexed hint chunks (e.g. "-hint-1", "-hint-2, "-hint-3", etc.) for your exercise chunk. For example:

<div id="exercisehints"></div>
<script type="text/javascript">loadSnippet('exercisehints')</script>

## Progressive Reveal

You might want users of your tutorials to see only one sub-topic at a time as they work through the material (this can be helpful to reduce distractions and maintain focus on the current task). If you specify the `progressive` option then all Level 3 headings (`###`) will be revealed progressively. For example:

<div id="tutorialyaml-progressive"></div>
<script type="text/javascript">loadSnippet('tutorialyaml-progressive', 'markdown')</script>

You can also specify this option on a per topic basis using the `data-progressive` attribute. For example, the following code enables progressive rendering for a single topic:

<div id="tutorialyaml-progressive-topic"></div>
<script type="text/javascript">loadSnippet('tutorialyaml-progressive-topic', 'markdown')</script>

You can also use `data-progressive=FALSE` to disable progressive rendering for an individual topic if the global `progressive` option is `TRUE`.

Progressive reveal provides an easy way to cue exercises one at a time: place each exercise under its own Level 3 header (`###`). This can be useful when a second exercises builds on the first, giving away the answer to the first.

Note that this feature is only available if you are using the [`learnr::tutorial`](formats.html#tutorial-format) R Markdown format (other custom formats may have their own internal mechanisms for progressive reveal).

### Exercise Skipping

When the `progressive` option is set to true, the tutorial will require students to complete any exercises in a sub-section before advancing to the next sub-section.

You may want to allow users of your tutorials to skip through exercises that they can't quite figure out. This might especially be true if you want users to be able to optionally see the next exercise even if they haven't completed the prior one.

If you specify the `allow_skip` option then students will be able to advance through a sub-section without completing the exercises. For example:

<div id="tutorialyaml-allowskip"></div>
<script type="text/javascript">loadSnippet('tutorialyaml-allowskip', 'markdown')</script>

You can also specify this option on a per sub-topic basis using the `data-allow-skip` attribute. For example, the following code enables exercise skipping within a single sub-topic:

<div id="tutorialyaml-allowskip-topic"></div>
<script type="text/javascript">loadSnippet('tutorialyaml-allowskip-topic', 'markdown')</script>

You can also use `data-allow-skip=FALSE` to disable exercise skipping rendering for an individual sub-topic if the global `allow-skip` option is `TRUE`.

## Exercise Checking

**learnr** provides allows full control over feedback provided to exercise submissions via `exercise.checker` in `tutorial_options()`. We'll eventually cover how to implement a custom `exercise.checker`, but for sake of demonstration, this section uses [**gradethis**](https://github.com/rstudio-education/gradethis)'s approach to exercise checking, which doesn't require knowledge of `exercise.checker` (**NOTE**: **gradethis** is still a work-in-progress. You may want to consider alternatives such as [**checkr**](https://github.com/dtkaplan/checkr)). To use **gradethis**'s approach to exercise checking inside of **learnr**, just call `gradethis_setup()` in a [setup](#Exercise_Setup) chunk, which will set `tutorial_options(exercise.checker = gradethis::grade_learnr)` (among other things).

### Checking results

Checking of exercise results may be done through a `*-check` chunk. With a **gradethis** setup, results can be graded with `grade_result()`, like so:

<div id="exercisecheck"></div>
<script type="text/javascript">loadSnippet('exercisecheck')</script>

When an exercise `*-check` chunk is provided, **learnr** provides an additional "Submit Answer" button, which allows users to experiment with various answers before formally submitting an answer:

<img src="images/exercise-result.png" width="100%" style="border: 1px solid #ddd; box-shadow:5px 5px 5px #eee;"/>

### Checking code

Checking of exercise *code* may be done through a `*-code-check` chunk. With a **gradethis** setup, if you supply a `*-solution` chunk and call `grade_code()` inside `*-code-check`, then you get detection of differences in the R syntax between the submitted exercise code and the solution code.

<div id="exercisecodecheck"></div>
<script type="text/javascript">loadSnippet('exercisecodecheck')</script>

<img src="images/exercise-code.png" width="100%" style="border: 1px solid #ddd; box-shadow:5px 5px 5px #eee;"/>

It's worth noting that, when a `*-code-check` chunk is supplied, the check is done *prior* to evaluation of the exercise submission, meaning that if the `*-code-check` chunk returns feedback, then that feedback is displayed, no exercise code is evaluated, and no result check is performed.


### Checking errors

In the event that an exercise submission generates an error, checking of the code (or its result, which is an error condition) may be done through either a `*-error-check` chunk or through the global `exercise.error.check.code` option. With a **gradethis** setup, `exercise.error.check.code` is set to `grade_code()`, meaning that if a `*-solution` chunk is provided, then more intelligent feedback is provided _for free_.

<img src="images/exercise-error.png" width="100%" style="border: 1px solid #ddd; box-shadow:5px 5px 5px #eee;"/>

To learn more about grading exercises with **gradethis**, see its grading demo ([`gradethis::gradethis_demo()`](https://github.com/rstudio-education/gradethis/blob/master/inst/tutorials/grading-demo/grading-demo.Rmd)).

### Custom Checking

If you need custom exercise checking logic that isn't already provided grading packages like [**gradethis**](https://github.com/rstudio-education/gradethis), then you may want to write your own `exercise.checker` function. This function is called on exercise submission whenever `*-check`, `*-code-check`, and/or `*-error-check` chunk(s) exist. When called, this function receives all the information that **learnr** knows about the exercise at the time of the checking, including the `user_code`, `solution_code`, `check_code`, exercise environments, and the `last_value` (if any). Checking can be performed at three different time points, so the values supplied can differ depending on the time point:

1. Before exercise evaluation, at this stage:
  * `check_code` contains `*-code-check` chunk code.
  * `envir_result`, `evaluate_result`, and `last_value` are all `NULL`.
2. During evaluation, when an error occurs:
  * `check_code` contains `*-error-check` chunk code.
  * `last_value` contains the error condition object.
3. After exercise evaluation, at this stage:
  * `check_code` contains `*-check` chunk code.
  * `last_value` contains the last printed value.

If, at any one of these stages, feedback should be provided, then `exercise.checker` should return an R list with, at the very least, a `correct` flag and feedback `message`:

<table>
<thead>
<tr class="header">
<th>Field</th>
<th>Description</th>
</tr>
</thead>
<tbody>
<tr class="odd">
<td><code>message</code></td>
<td>Feedback message. Can be a plain character vector or can HTML produced via the <a href="https://cran.r-project.org/web/packages/htmltools/index.html">htmltools</a> package.</td>
</tr>
<tr class="even">
<td><code>correct</code></td>
<td>TRUE/FALSE logical value indicating whether the submitted answer was correct.</td>
</tr>
<tr class="odd">
<td><code>type</code></td>
<td>Feedback type (visual presentation style). Can be "auto", "success", "info", "warning", "error", or "custom". Note that "custom" implies that the "message" field is custom HTML rather than a character vector.</td>
</tr>
<tr class="even">
<td><code>location</code></td>
<td>Location for feedback ("append", "prepend", or "replace").</td>
</tr>
</tbody>
</table>

Below is a rough sketch of how one might implement an `exercise.checker` function. Notice how the presence of `envir_result` may be used to determine which type of check is being done (i.e., code or result check). 

<div id="customchecker"></div>
<script type="text/javascript">loadSnippet('customchecker')</script>

See the table below for a full description of all the arguments supplied to `exercise.checker`.

<table>
<thead>
<tr class="header">
<th>Argument</th>
<th>Description</th>
</tr>
</thead>
<tbody>
<tr class="odd">
<td><code>label</code></td>
<td>Label for exercise chunk.</td>
</tr>
<tr class="even">
<td><code>user_code</code></td>
<td>R code submitted by the user.</td>
</tr>
<tr class="odd">
<td><code>solution_code</code></td>
<td>Code provided within the "-solution" chunk for the exercise.</td>
</tr>
<tr class="even">
<td><code>check_code</code></td>
<td>Code provided within the "-check" chunk for the exercise.</td>
</tr>
<tr class="odd">
<td><code>envir_result</code></td>
<td>The R environment <b>after</b> the execution of the chunk.</td>
</tr>
<tr class="even">
<td><code>evaluate_result</code></td>
<td>The return value from the <a href="https://www.rdocumentation.org/packages/evaluate/topics/evaluate"><code>evaluate::evaluate</code></a> function.</td>
</tr>
<tr class="odd">
<td><code>envir_prep</code></td>
<td>A copy of the R environment <b>before</b> the execution of the chunk.</td>
</tr>
<tr class="even">
<td><code>last_value</code></td>
<td>The last value of the evaluated chunk.</td>
</tr>
<tr class="odd">
<td><code>...</code></td>
<td>Unused (include for compatibility with parameters to be added in the future)</td>
</tr>
</tbody>
</table>




## Exercise Caption

By default exercises are displayed with caption of "Code". However, in some cases you may want either a custom per-chunk caption or a generic caption with a different connotation (e.g. "Exercise" or "Sandbox"). For example:

<div id="exercisecaption"></div>
<script type="text/javascript">loadSnippet('exercisecaption')</script>

## Code Assistance

By default, code completions are automatically provided as users type within the exercise editor:

<img src="images/exercise-completion.png" width=749 height=213>

You can optionally disable code completion (either globally or on a per-chunk basis) using the `exercise.completion` option. For example, the following illustrates turning off completions globally then enabling them for an individual chunk:

<div id="exercisecompletion"></div>
<script type="text/javascript">loadSnippet('exercisecompletion')</script>

Similarily, simple code diagnostics can also be provided, to inform users of errors in the R code written in exercise editors. Diagnostics can be controlled (either globally or on a per-chunk basis) using the `exercise.diagnostics` option.

## Editor Size

By default, the size of the exercise editor provided to users will match the number of lines in your code chunk (with a minimum of 2 lines). If the user adds additional lines in the course of editing the editor will grow vertically up to 15 lines, after which it will display a scrollbar.

You can also specify a number of lines explicitly using the `exercise.lines` chunk option (this can be done on a per-chunk or global basis). For example, the following chunk specifies that the exercise code editor should be 15 lines high:

<div id="exerciselines"></div>
<script type="text/javascript">loadSnippet('exerciselines')</script>

## Time Limits

To mediate the problem of code which takes longer than expected to run you can specify the `exercise.timelimit` chunk option or alternatively the global `tutorial.exercise.timelimit` option.

The following demonstrates setting a 10 second time limit as a global option, document level option, and chunk level option:

<div id="exercisetimelimit"></div>
<script type="text/javascript">loadSnippet('exercisetimelimit')</script>

Since tutorials are a highly interactive format you should in general be designing exercises that take no longer than 5 or 10 seconds to execute. Correspondingly, the default value for `tutorial.exercise.timelimit` if not otherwise specified is 30 seconds.
